---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
NAME = "ISREAL UFUMAKA"
COLLABORATORS = "2204569"
```

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=100)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300

def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    # YOUR CODE HERE
    return sample.mean()
    raise NotImplementedError()
```

Output:
```
253.76
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

---
## Question 2: 

#### Majority Voting
Given a list = `[1,2,3,1,2,1,1,1,3,2,2,1]`. Write a function to find the majority voting.

```python
from collections import Counter
```

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
a = np.array([1,2,3,1,2,1,1,1,3,2,2,1])

def my_majorityvote():
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    c = Counter(a)
    return c.most_common(1)[0][0]
    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 1, rel_tol = 0.05)
```

---
## Question 3: 

#### Decision Tree
USe the IRIS dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely IRIS Setosa, IRIS Versicolour, and IRIS Virginica, with the following attributes: sepal length, sepal width, petal length, and petal width. Your aim is to train a decision tree on training data and return the test accuracy.

```python
NAME = "ISREAL UFUMAKA"
COLLABORATORS = "2204569"
```

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    # YOUR CODE HERE
     # Create a decision tree classifier
    clf = DecisionTreeClassifier()

    # Train the classifier using the training data
    clf.fit(X_train, y_train)

    # Predict labels for the test data
    y_pred = clf.predict(X_test)

    # Calculate the test accuracy
    accuracy = accuracy_score(y_test, y_pred)
    return accuracy
    raise NotImplementedError()
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 0.94736, rel_tol = 0.05)
```

```python

```

```python

```