---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 



```python
NAME = "ISREAL UFUMAKA"
COLLABORATORS = "2204569"
```

```python
import numpy as np
from sklearn import datasets, linear_model
from sklearn.metrics import mean_squared_error, r2_score

# Load the diabetes dataset
diabetes_X, diabetes_y = datasets.load_diabetes(return_X_y=True)

# Use only one feature
diabetes_X = diabetes_X[:, np.newaxis, 2]

# Split the data into training/testing sets
diabetes_X_train = diabetes_X[:-20]
diabetes_X_test = diabetes_X[-20:]

# Split the targets into training/testing sets
diabetes_y_train = diabetes_y[:-20]
diabetes_y_test = diabetes_y[-20:]
```

---
## Question 1: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `mean squared error`.  

<!-- grade_id: my_reg_model_mse_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `mean_squared_error` on diabetes_X_test. 
    Please use LinearRegression model to train a regressor. """
    # YOUR CODE HERE
    reg = linear_model.LinearRegression()

    # Train the model using the training data
    reg.fit(diabetes_X_train, diabetes_y_train)

    # Predict labels for the test data
    y_pred = reg.predict(diabetes_X_test)

    # Calculate the mean squared error
    mse = mean_squared_error(diabetes_y_test, y_pred)

    return mse
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_mse_test, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), 
                    2548.07, rel_tol = 0.05)
```

---
## Question 2: 

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `r2` (r sqaured value).  

<!-- grade_id: my_reg_model_r2_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `r2_score` on diabetes_X_test data . 
    Please use LinearRegression model to train a regressor. """
    # YOUR CODE HERE
    # Create a linear regression model
    reg = linear_model.LinearRegression()
    # Train the model using the training data
    reg.fit(diabetes_X_train, diabetes_y_train)

    # Predict labels for the test data
    y_pred = reg.predict(diabetes_X_test)

    # Calculate the R^2 score
    r2 = r2_score(diabetes_y_test, y_pred)

    return r2
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_r2_test, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
                                    diabetes_y_test), 0.472575, rel_tol = 0.05)
```

```python

```