---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and return log2 of the array in the list format. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0.0, 2.0, 3.0]` 

```python
NAME = "ISREAL UFUMAKA"
COLLABORATORS = "2204569"
```

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myLog(mylist):
    """Write a function to pass a list of numbers and return log2 of the array. """
    # YOUR CODE HERE
    mylist = np.array(mylist)
    # print(mylist)
    return list(np.log2(mylist))

    raise NotImplementedError()
```

```python
in_array = [1, 4, 8]
myLog(in_array)
```

Output:
```
[0.0, 2.0, 3.0]
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case 1
my = [1, 3, 5, 10**8]
assert myLog(my) == [0.0, 1.584962500721156, 2.321928094887362, 26.575424759098897]

# Test case 2
my = [1, 100, 1000]
assert myLog(my) == [0.0, 6.643856189774724, 9.965784284662087]
```

---

## Question 2:

The Fibonacci numbers are the numbers in the following integer sequence.
`[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, ...]`
In mathematical terms, the sequence $F_n$ of Fibonacci numbers is defined by the recurrence relation 

$$
F_n = F_{n-1} + F_{n-2}
$$
with seed values  
$$
 F_0 = 0 \& F_1 = 1.
$$


Your task here is to write a code for printing $n^{th}$ Fibonacci number
`Example`: Input: `myfibonacci(9)` and Output should be `34`. Remember your index starts from 0.

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def myfibonacci(n):
    """Function for nth Fibonacci number"""
    # YOUR CODE HERE
    if n == 0:
      return 0
    elif n == 1:
      return 1
    
    return myfibonacci(n-1) + myfibonacci(n-2)

    raise NotImplementedError()
```

```python
myfibonacci(9)
```

Output:
```
34
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct fibonacci number"""
assert myfibonacci(0) == 0
assert myfibonacci(1) == 1
assert myfibonacci(2) == 1
assert myfibonacci(3) == 2
assert myfibonacci(4) == 3
assert myfibonacci(5) == 5
assert myfibonacci(6) == 8
assert myfibonacci(7) == 13
assert myfibonacci(8) == 21
assert myfibonacci(9) == 34
```

---
## Question 3: 

 Write a program, which accept a list with `nan`s in it and count the maximum number of consecutive repeated `nan`.
 Example: In the `a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]` and answer is 2.

<!-- grade_id: max_nan_ans, points: 0, grade: false, solution: true -->
```python
from numpy import nan
# YOUR CODE HERE
def max_rep_nan(mylist):
  # creating a counter and an empty list to hold counts
  count = 0
  counter_holder = []

  # checking if an empty list is passed and returning 0 if no empty
  if len(mylist) == 0:
    return 0
  # looping through my list to check which item is nan and keeping count
  for item in mylist:
    # print(item)
    if np.isnan(item):
      count += 1
    else:
      # resetting counter if next index doesn't contain nan
      counter_holder.append(count)
      count = 0
  # print(counter_holder)
  return max(counter_holder)
  raise NotImplementedError()
```

```python
a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]
max_rep_nan(a)
```

Output:
```
2
```

<!-- grade_id: max_rep_test, points: 0, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 3

a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 5

a = [nan, 10, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 1

a = []
assert max_rep_nan(a) == 0
```

## Pandas Data Frame 


```python
from sklearn import datasets
import pandas as pd
```

```python
# import some data to play with
iris = datasets.load_iris()
X = pd.DataFrame(iris.data[:, :])  # we are taking all features.
y = iris.target
print(X)
print(y)
```

Output:
```
0    1    2    3
0    5.1  3.5  1.4  0.2
1    4.9  3.0  1.4  0.2
2    4.7  3.2  1.3  0.2
3    4.6  3.1  1.5  0.2
4    5.0  3.6  1.4  0.2
..   ...  ...  ...  ...
145  6.7  3.0  5.2  2.3
146  6.3  2.5  5.0  1.9
147  6.5  3.0  5.2  2.0
148  6.2  3.4  5.4  2.3
149  5.9  3.0  5.1  1.8

[150 rows x 4 columns]
[0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0
 0 0 0 0 0 0 0 0 0 0 0 0 0 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1
 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 1 2 2 2 2 2 2 2 2 2 2 2
 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2 2
 2 2]
```

---
## Question 4:

Write a function to find the max number in the each column in the IRIS dataset and return the output as a list,

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def my_max_min(X):
    """Write a function to find the max number in the each column and return 
    as a list."""
    # YOUR CODE HERE
    return list(X.max())
    raise NotImplementedError()
```

```python

```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_max_min(X) == [7.9, 4.4, 6.9, 2.5]
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of IRIS dataset (last column only). `Hint`: You can use NumPy aggregate functions. Create a list and append value to it and return the output. 

```python
# from decimal import Decimal
```

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of IRIS dataset (last column only)."""
    # YOUR CODE HERE
    my_stat = [X[3].mean(), X[3].std()]
    print(f"Returned Values {my_stat}")
    return my_stat
    raise NotImplementedError()
```

<!-- grade_id: stats_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
assert stats(X) == [1.1993333333333336, 0.7622376689603465]
```

Output:
```
Returned Values [1.199333333333334, 0.7622376689603465]
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_11716\2933274922.py in <module>
      1 """Check your output"""
      2 # Test case
----> 3 assert stats(X) == [1.1993333333333336, 0.7622376689603465]

AssertionError:
```

The error you are seeing is likely caused by the fact that the assert statement is comparing two lists of floating-point numbers, which may not be exactly equal due to the way they are represented in memory.

To fix this issue, you can use the assert math.isclose statement to check if the values are close to the expected values within a certain tolerance, specified by the rel_tol argument.

```python
import math

# Test case
assert math.isclose(stats(X)[0], 1.1993333333333336, rel_tol=1e-9)
assert math.isclose(stats(X)[1], 0.7622376689603465, rel_tol=1e-9)

```

Output:
```
Returned Values [1.199333333333334, 0.7622376689603465]
Returned Values [1.199333333333334, 0.7622376689603465]
```

```python

```