---
# Welcome to CE880
### This is your week-3 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

We have already loaded `iris` dataset for you. Now, write a function to calculate the `median` on the 1st column of the dataset. 
`Hint`: To find the median, you first need to sort the values in your sample. But, before that we have sliced, the first column of `iris` data into a variable name `sample` and data type is a list. Then pass `sample` list as a parameter in you `my_median` function.

The `median` of a sample of numeric data is the value that lies in the middle when we sort the data. The data may be sorted in ascending or descending order, the median remains the same. Please also consider what to do if there is no middle value.

Please don't use `np.median` function. Write your own custom function

```python
NAME = "ISREAL UFUMAKA"
COLLABORATORS = "2204569"
```

<!-- grade_id: my_median_ans, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets
import statistics
iris = datasets.load_iris()
sample = list(iris.data[:,0])

def my_median(sample):
    """Function for calculating median"""
    return statistics.median(sample)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

<!-- grade_id: my_median_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_median(sample) == 5.8
```

---

## Question 2:

`Please read me before attemping the question`

A quartile is a type of quantile. The first quartile `(Q1)`, is defined as the middle number between the smallest number and the median of the data set, the second quartile `(Q2)` – median of the given data set while the third quartile `(Q3)`, is the middle number between the median and the largest value of the data set.

Algorithm to find Quartiles :
Quartiles are calculated by the help of the median. If the number of entries is an even number i.e. of the form 2n, then, first quartile `(Q1)` is equal to the median of the `n` smallest entries and the third quartile `(Q3)` is equal to the median of the `n` largest entries.

If the number of entries is an odd number i.e. of the form `(2n + 1)`, then

* the first quartile (Q1) is equal to the median of the n smallest entries
* the third quartile (Q3) is equal to the median of the n largest entries
* the second quartile(Q2) is the same as the ordinary median.

Example: 
input = [75, 69, 56, 46, 47, 79, 92, 97, 89, 88, 36, 96, 105, 32, 116, 101, 79, 93, 91, 112]
After sorting the above data set:
[32, 36, 46, 47, 56, 69, 75, 79, 79, 88, 89, 91, 92, 93, 96, 97, 101, 105, 112, 116]
Here the total number of terms is 20.

The second quartile `(Q2)` or the median of the above data is `(88 + 89) / 2 = 88.5`
The first quartile `(Q1)` is median of first `n` i.e. 10 terms (or `n` i.e. 10 smallest values) = 62.5
The third quartile `(Q3)` is the median of `n` i.e. 10 largest values (or last `n` i.e. 10 values) = 96.5
Then, `IQR = Q3 – Q1 = 96.5 – 62.5 = 34.0`


Now, your job is to write a function to Caluate Interquartile Range on the second column of `iris` dataset, which you have loaded in the question 1.
`hint`: We have created a new variable `sample2` and sliced the second column of the `iris` data. Please explore `np.percentile` function

```python
test_sample = [75, 69, 56, 46, 47, 79, 92, 97, 89, 88, 36, 96, 105, 32, 116, 101, 79, 93, 91, 112]
test_sample.sort()
```

<!-- grade_id: my_quartile_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
sample2 = list(iris.data[:,1])

def my_quartile(sample2):
    """Caluate Interquartile Range of the given list."""
    # YOUR CODE HERE
    num = len(sample2)
    sample2.sort()
    Q2 = my_median(sample2)
    Q1 = my_median(sample2[:(num//2)])
    Q3 = my_median(sample2[(num//2):])
    
    return Q3 - Q1
    raise NotImplementedError()
```

```python
#test checks
Q1 = my_median(test_sample)
Q2 = my_median(test_sample[:10])
Q3 = my_median(test_sample[10:])
Q1, Q2, Q3
```

Output:
```
(88.5, 62.5, 96.5)
```

```python
#test checks
my_quartile(test_sample)
```

Output:
```
34.0
```

<!-- grade_id: my_quartile_test, points: 2, grade: true, solution: false -->
```python
"""Check that correct answers here"""
import math
assert math.isclose(my_quartile(sample2), 0.5, rel_tol=0.05)
```

---
## Question 3: 

Write a function to calculate the `standard deviation` on the third column of `iris` dataset. We have created a new variable `sample3` and sliced the third column of the iris data. Please create a function that first calculate the variance then calculate `standard deviation`. Please do not use inbuilt `numpy.var` and `numpy.std` functions. Write your own codes. 

<!-- grade_id: my_std_ans, points: 0, grade: false, solution: true -->
```python
import math
sample3 = list(iris.data[:,2])
def my_std(sample3, ddof=0):
    """Write a program to calculate the standard deviation of 3rd column of iris dataset."""
    
    # YOUR CODE HERE
    mean = my_mean(sample3)
    dev = []
    for sample in sample3:
      # print(sample)
      dev.append((sample - mean) ** 2)
    
    var = sum(dev)/(len(sample3)-1)
    std = math.sqrt(var)
    return std
    raise NotImplementedError()
```

```python
mean = my_mean(sample3)
mean
```

Output:
```
3.7580000000000027
```

```python

```

<!-- grade_id: my_std_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_std(sample3), 1.759, rel_tol=0.05)
```

---
## Question 4:

Write a program to find the mean of the 4th column of `iris` dataset. 

<!-- grade_id: my_mean_ans, points: 0, grade: false, solution: true -->
```python
sample4 = list(iris.data[:,3])

def my_mean(sample4):
    """write a function to find the mean of the 4th column of iris data"""
    # YOUR CODE HERE
    return sum(sample4)/len(sample4)
    raise NotImplementedError()
```

<!-- grade_id: my_mean_test, points: 1, grade: true, solution: false -->
```python
"""Test you solution"""
assert math.isclose(my_mean(sample4), 1.19, rel_tol=0.05)
```

## Question 5

Write a Pandas program to calculate the correlation between column 1 and column 3 of `iris` dataset. `Hint` use `from scipy import stats` to find a Pearson correlation coefficient. 

<!-- grade_id: iris_correlation_ans, points: 0, grade: false, solution: true -->
```python
import itertools
from scipy import stats
sample1 = list(iris.data[:,0])
sample2 = list(iris.data[:,2])
def iris_correlation(sample1, sample2):
    """calculate the correlation"""
    # YOUR CODE HERE
    pers = stats.pearsonr(sample1, sample2)
    return pers[0]
    raise NotImplementedError()
```

```python
stats.pearsonr(sample1, sample2)
```

Output:
```
(0.8717537758865831, 1.0386674194498099e-47)
```

```python
"""Test you solution"""
assert math.isclose(iris_correlation(sample1, sample2), 0.87, rel_tol=0.05)
```

```python

```